<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Einführung Machine Learning
### Sommersemester 2026
Prof. Dr. Heiner Giefers

## Aufgabe zur Logistischen Regression

### Einleitung und Datensatzbeschreibung

Der vorliegende Datensatz untersucht den schulischen Erfolg von Schüler:innen in der Sekundarstufe an zwei portugiesischen Schulen. Er wurde im Rahmen einer Studie von Cortez und Silva (2008) erstellt und umfasst neben den Schulnoten auch demografische, soziale und schulbezogene Merkmale. Die Daten wurden sowohl aus offiziellen Schulberichten als auch über Fragebögen erhoben.

Es existieren zwei separate Datensätze, die sich auf die Leistungen in zwei unterschiedlichen Fächern beziehen: Mathematik (`student-mat.csv`) und Portugiesisch (`student-por.csv`). Beide Datensätze enthalten dieselbe Struktur mit 33 Attributen. Die Zielvariable `G3` steht für die Abschlussnote am Ende des Schuljahres und ist die zentrale Variable für Klassifikations- und Regressionsaufgaben.

Ein wichtiger Aspekt: `G3` ist stark mit den Zwischennoten `G1` (Note im 1. Trimester) und `G2` (Note im 2. Trimester) korreliert. Dies ist naheliegend, da alle drei Noten den schulischen Fortschritt desselben Fachs im Verlauf des Schuljahres abbilden. Obwohl eine Vorhersage von `G3` ohne die Kenntnis von `G1` und `G2` schwieriger ist, ist sie für praktische Anwendungsfälle relevanter, da man die Abschlussleistung möglichst frühzeitig prognostizieren möchte.

Der Datensatz eignet sich gut für Machine-Learning-Aufgaben wie Klassifikation (z. B. Bestehen vs. Nicht-Bestehen) und Regression (z. B. exakte Notenprognose) und stellt durch seine realweltliche Herkunft ein realistisches Anwendungsszenario dar.

```python
import pandas as pd
import os
import urllib.request

url = "https://raw.githubusercontent.com/fhswf/datasets/main/student-por.csv"

dfile = "./student.csv"

if not os.path.isfile(dfile):
    urllib.request.urlretrieve(url, dfile)
```

```python
import pandas as pd
df = pd.read_csv(dfile)
df.head(10)
```

Mit `df.info()` und `df.describe()` erhalten wir einige Informationen über den Datensatz.

```python
df.info()
```

```python
df.describe()
```

### Merkmale des Datensatzes

Der Datensatz enthält soziodemografische und schulbezogene Informationen über Schüler, die zur Vorhersage schulischer Leistungen (G3) verwendet werden können. Im Folgenden sind alle verfügbaren Merkmale aufgeführt:

- **school**: Schule des Schülers (z. B. 'GP' = Gabriel Pereira, 'MS' = Mousinho da Silveira)
- **sex**: Geschlecht (M = männlich, F = weiblich)
- **age**: Alter in Jahren
- **address**: Wohnort (U = urban, R = ländlich)
- **famsize**: Familiengröße (LE3 = ≤ 3 Personen, GT3 = > 3 Personen)
- **Pstatus**: Zusammenleben der Eltern (T = zusammen, A = getrennt)
- **Medu**: Bildungsniveau der Mutter (0 = keine, 4 = Hochschulabschluss)
- **Fedu**: Bildungsniveau des Vaters (0 = keine, 4 = Hochschulabschluss)
- **Mjob**: Beruf der Mutter (z. B. teacher, health, services, at_home, other)
- **Fjob**: Beruf des Vaters (wie oben)
- **reason**: Grund für Wahl der Schule (z. B. home, reputation, course, other)
- **guardian**: Erziehungsberechtigter (mother, father, other)
- **traveltime**: Schulweg in Minuten (1 = <15min, 4 = >1h)
- **studytime**: Wöchentliche Lernzeit in Stunden (1 = <2h, 4 = >10h)
- **failures**: Anzahl bisher nicht bestandener Schuljahre
- **schoolsup**: Zusätzliche schulische Unterstützung (yes/no)
- **famsup**: Familiäre Lernunterstützung (yes/no)
- **paid**: Teilnahme an bezahlten Zusatzkursen (yes/no)
- **activities**: Teilnahme an Freizeitaktivitäten (yes/no)
- **nursery**: Besuch einer Vorschule (yes/no)
- **higher**: Wunsch nach höherer Bildung (yes/no)
- **internet**: Internetzugang zu Hause (yes/no)
- **romantic**: In einer Beziehung (yes/no)
- **famrel**: Qualität der Familienbeziehungen (1 = sehr schlecht, 5 = sehr gut)
- **freetime**: Freizeit nach der Schule (1 = sehr wenig, 5 = sehr viel)
- **goout**: Ausgehverhalten mit Freunden (1 = sehr selten, 5 = sehr häufig)
- **Dalc**: Alkoholkonsum werktags (1 = sehr niedrig, 5 = sehr hoch)
- **Walc**: Alkoholkonsum am Wochenende (1 = sehr niedrig, 5 = sehr hoch)
- **health**: Eigene Gesundheitseinschätzung (1 = sehr schlecht, 5 = sehr gut)
- **absences**: Anzahl der Fehltage in der Schule
- **G1**: Note im 1. Trimester (0–20)
- **G2**: Note im 2. Trimester (0–20)
- **G3**: Endnote (0–20)

### Zielvariable: Hat der Schüler bestanden?

In dieser Zelle wird die Zielvariable `bestanden` eingeführt. Ein Schüler gilt als **bestanden**, wenn die Abschlussnote `G3` größer als 10 ist. Daraus wird eine binäre Klassifikationsvariable abgeleitet (1 = bestanden, 0 = nicht bestanden). Anschließend wird die Verteilung dieser Zielklasse ausgegeben.

Diese Zielvariable bildet die Grundlage für die Klassifikationsaufgabe im weiteren Verlauf.

**Hinweis:** Falls `G3` fehlende Werte enthält, wäre eine Vorabprüfung sinnvoll.

```python
# Ziel: Hat der Schüler bestanden? Wir definieren: bestanden = G3 > 10
df['bestanden'] = (df['G3'] > 10).astype(int)
df['bestanden'].value_counts()
```

### One-Hot-Encoding: Umwandlung kategorialer Merkmale

In dieser Zelle werden alle **nicht-numerischen Spalten** des Datensatzes in numerische Form überführt, indem **One-Hot-Encoding** angewendet wird. Dabei wird jede kategoriale Variable in binäre Spalten umgewandelt. Mit `drop_first=True` wird dabei eine Dummy-Falle vermieden, indem jeweils eine Kategorie pro Merkmal weggelassen wird.

Die Zielvariable `y` wird separat gespeichert, und `X` enthält die daraus vorbereiteten Eingabe-Features. Die ursprünglichen Noten `G1`, `G2`, `G3` sowie die Zielvariable `bestanden` werden aus `X` entfernt, da sie entweder Ziel oder stark korreliert mit dem Ziel sind.

```python
# Nicht-numerische Spalten umwandeln (One-Hot-Encoding)
df_encoded = pd.get_dummies(df.drop(columns=['G1', 'G2', 'G3','bestanden']), drop_first=True)
X = df_encoded

y = df['bestanden']
```

### Aufteilung in Trainings- und Testdaten

Hier wird der vorbereitete Datensatz in **Trainings- und Testdaten** aufgeteilt. Mit `train_test_split` aus `sklearn.model_selection` werden:

- 70 % der Daten für das Training (`X_train`, `y_train`)
- 30 % für das Testen (`X_test`, `y_test`)

verwendet. Der Parameter `random_state=42` stellt sicher, dass die Aufteilung reproduzierbar ist.

Diese Trennung ist essenziell, um ein Modell unabhängig zu trainieren und anschließend seine Generalisierungsfähigkeit zu bewerten.

```python
from sklearn.model_selection import train_test_split

X_train, X_test, y_train, y_test = train_test_split(
    X, y, test_size=0.3, random_state=42
)
```

### Modelltraining mit logistischer Regression

Hier wird ein **logistisches Regressionsmodell** aus `sklearn.linear_model` erstellt und mit den Trainingsdaten trainiert:

- `LogisticRegression(max_iter=1000)` legt fest, dass die Optimierung bis zu 1000 Iterationen erlaubt ist (Standard ist 100, was bei vielen Features evtl. nicht ausreicht).
- `model.fit(X_train, y_train)` führt das Training durch.

Die logistische Regression ist ein einfaches, gut interpretierbares Modell zur **binären Klassifikation** – in diesem Fall: *Hat der Schüler bestanden (1) oder nicht (0)?*

```python
from sklearn.linear_model import LogisticRegression

model = LogisticRegression(max_iter=1000)
model.fit(X_train, y_train)
```

### Modellbewertung: Konfusionsmatrix und Klassifikationsbericht

Nach dem Training wird das Modell auf den Testdaten evaluiert:

- `model.predict(X_test)` erzeugt Vorhersagen für die Testmenge.
- `confusion_matrix(y_test, y_pred)` zeigt, wie viele Schüler korrekt/nicht korrekt als bestanden/nicht bestanden klassifiziert wurden.
- `classification_report(y_test, y_pred)` liefert präzise Kennzahlen wie **Precision**, **Recall** und **F1-Score** für beide Klassen.

Diese Auswertung hilft zu beurteilen, wie gut das Modell im Klassifikationsproblem "Bestanden oder nicht?" abschneidet.

```python
from sklearn.metrics import classification_report, confusion_matrix

y_pred = model.predict(X_test)
print(confusion_matrix(y_test, y_pred))
print(classification_report(y_test, y_pred))
```

### Wichtigste Merkmale: Einfluss auf das Bestehen

Mit `model.coef_` lassen sich die **gelernten Gewichtungen** der logistischen Regression auslesen. Diese geben an, wie stark ein Merkmal (positiv oder negativ) die Wahrscheinlichkeit beeinflusst, dass ein Schüler besteht.

- Die `Series` wird nach den höchsten Koeffizienten sortiert.
- Die Top 10 Merkmale mit **positivem Einfluss** auf das Bestehen werden angezeigt.

Dies ermöglicht eine erste Interpretation, welche Faktoren laut Modell besonders relevant sind.

```python
coeffs = pd.Series(model.coef_[0], index=X.columns)
coeffs.sort_values(ascending=False).head(10)
```

### ROC-Kurve und AUC: Modellgüte visualisieren

Die **ROC-Kurve** (Receiver Operating Characteristic) zeigt die Leistung des Klassifikators bei verschiedenen Schwellenwerten. Je näher die Kurve an der linken oberen Ecke verläuft, desto besser das Modell.

- Die **True Positive Rate (TPR)** wird gegen die **False Positive Rate (FPR)** aufgetragen.
- Der **AUC-Wert** (Area Under Curve) quantifiziert die Modellgüte (1 = perfekt, 0.5 = Zufall).
- `predict_proba()` liefert die geschätzte Wahrscheinlichkeit für Klasse 1 („bestanden“).

Eine gute Klassifikation zeigt sich in einem hohen AUC-Wert (> 0.8 ist i.d.R. stark).

```python
import matplotlib.pyplot as plt
from sklearn.metrics import roc_curve, roc_auc_score

probs = model.predict_proba(X_test)[:, 1]
fpr, tpr, _ = roc_curve(y_test, probs)
auc = roc_auc_score(y_test, probs)

plt.plot(fpr, tpr, label=f'AUC = {auc:.2f}')
plt.xlabel('False Positive Rate')
plt.ylabel('True Positive Rate')
plt.title('ROC-Kurve')
plt.legend()
plt.grid(True)
plt.show()
```

**Aufgabe:**  
Unter der URL [https://raw.githubusercontent.com/fhswf/datasets/main/student-mat.csv](https://raw.githubusercontent.com/fhswf/datasets/main/student-mat.csv) finden Sie einen sehr ähnlichen Datensatz, für das Fach Mathematik
Welche *Classification Accuracy* erreichen Sie bei der Vorhersage, ob ein Schüler bestanden hat (d.h. `G3 > 10`)?

Gehen Sie dabei analog zur bisherigen Analyse vor:

1. Laden Sie den Datensatz.
2. Definieren Sie die Zielvariable `bestanden`.
3. Wenden Sie One-Hot-Encoding auf kategoriale Merkmale an.
4. Trainieren Sie ein logistisches Regressionsmodell.
5. Bewerten Sie das Modell mit der Classification Accuracy.


<!-- grade_id: cell-f132d0c0e1b85258, points: 0, grade: true, solution: true -->
```python
url = "https://raw.githubusercontent.com/fhswf/datasets/main/student-mat.csv"
# YOUR CODE HERE
raise NotImplementedError()
```

**Aufgabe:**

Verwenden Sie die oben beschrieben Techniken, um Entscheidungsgrenzen für einen zufällig erzeugten Datensatz zu berechnen.

```python
from sklearn.datasets import make_blobs
import numpy as np
import matplotlib.pyplot as plt


# generating two-class dataset
X, y = make_blobs(n_samples=100, centers=2, n_features=2, center_box = (-5, 5))


plt.scatter(X[:, 0], X[:, 1], c=y, cmap=plt.cm.Spectral)
```

**(a)** Teilen Sie den Datensatz auf (70% Training, 30% Test):
- `X_train`: training dataset
- `X_test`: test dataset
- `y_train`: training labels
- `y_test`: test labels

<!-- grade_id: cell-e7c0c8cf666d3b5c, points: 0, grade: false, solution: true -->
```python
X_train, X_test, y_train, y_test = [None]*4
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8818e07d53ed936b, points: 1, grade: true, solution: false -->
```python
# Test Cell
#----------

assert np.vstack((X_train, X_test)) in X and np.hstack((y_train, y_test)) in y
assert y_train.size/y.size == 0.7
```

**(b)** Verwenden Sie die *sklearn*-Klasse `LogisticRegression` um ein Modell für den Datensatz zu bilden. Trainieren Sie das Modell mit den oben festgelegten Trainingsdaten.

<!-- grade_id: cell-12524a6611e42cf3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c2da293e49378c98, points: 1, grade: true, solution: false -->
```python
# Test Cell
#----------

assert type(logreg) == LogisticRegression
assert logreg.intercept_, 'Trainieren Sie das Modell mit den Daten!'
```

Wir visualisieren nun den Datensatz um darzustellen, wie gut unser Modell klassifiziert:

```python
# Plotting decision regions
x_min, x_max = X[:, 0].min() - 1, X[:, 0].max() + 1
y_min, y_max = X[:, 1].min() - 1, X[:, 1].max() + 1
xx, yy = np.meshgrid(np.arange(x_min, x_max, 0.01),np.arange(y_min, y_max, 0.01))

Z = logreg.predict(np.c_[xx.ravel(), yy.ravel()])
Z = Z.reshape(xx.shape)

plt.contourf(xx, yy, Z, alpha=0.4, cmap=plt.cm.Spectral)
plt.scatter(X[:, 0], X[:, 1], c=y, cmap=plt.cm.Spectral)
plt.xlabel(r"$\Theta_0$", fontsize=14)
plt.ylabel(r"$\Theta_1$", fontsize=14)
plt.show()
```

**(c)** Testen Sie die Vorhersagegenauigkeit (*accuracy*) des Modells mit den Testdaten:

<!-- grade_id: cell-6f17cc003e3def17, points: 0, grade: false, solution: true -->
```python
y_pred = None
acc_test = None

# YOUR CODE HERE
raise NotImplementedError()

print("accuracy: %.2f%%" % acc_test)
```

<!-- grade_id: cell-35985385625486b8, points: 1, grade: true, solution: false -->
```python
# Test Cell
#----------

assert y_pred.shape == y_test.shape
### BEGIN HIDDEN TEST
assert acc_test == 100-np.sum(np.abs(y_pred-y_test))*100/len(y_pred)
### END HIDDEN TEST
```

### Multiklassen-Klassifikation 

Die logistische Regression liefert uns Ergebnisse für binäre Zielvariablen.
Oftmals wollen wir aber mehr als 2 Klassen unterscheiden.

Eine Möglichkeit, um Multiklassen-Klassifikation mit logistischen Regression umzusetzen, ist die sogenannte *One-vs-all Klassifikation*.
Dabei werden für `n` Klassen `n` separate, binäre Klassifikationsprobleme definiert, bei denen jeweils nur die betrachtete Klasse den Zielwert `1` zugeteilt bekommt, und für alle anderen Klassen der Zielwert `0` angenommen wird.

Scikit-Learn unterstüzt Multiklassen-Klassifikation in der Klasse `LogisticRegression` über den Parameter `multi_class`.
Setzt man : `multi_class="ovr"` führt die Funktion `fit` je eine logistische Regression für jedes Label nach dem *one-vs-all* (oder auch *one-vs-rest*, ovr) Prinzip aus.

In der folgenden Code-Zelle erzeugen wir 3 Punktwolken.
Alle Punkte einer "Wolke" sollen zu einer bestimmten Klasse gehören.

```python
from sklearn.datasets import make_blobs
import matplotlib.pyplot as plt
from pandas import DataFrame
# generate 2d classification dataset
X, y = make_blobs(n_samples=100, centers=3, n_features=2, random_state=10, cluster_std=2.5)
# scatter plot, dots colored by class value
df = DataFrame(dict(x=X[:,0], y=X[:,1], label=y))
colors = {0:'red', 1:'blue', 2:'green'}
markers = {0:'o', 1:'x', 2:'^'}
fig, ax = plt.subplots()
grouped = df.groupby('label')
for key, group in grouped:
    group.plot(ax=ax, kind='scatter', x='x', y='y', label=key, marker=markers[key], color=colors[key])
    
plt.legend(loc='upper right', prop={'size': 12})
plt.savefig("LogistischeRegression20.png",transparent=True, dpi=300)
plt.show()
```

Nun wenden wir ein logistisches Regressionsmodell auf die Datenbasis an.

```python
from sklearn.preprocessing import MinMaxScaler
from sklearn.linear_model import LogisticRegression

X_train, X_test, y_train, y_test = train_test_split(X,y,test_size=0.3, random_state=0)

model = LogisticRegression(solver='lbfgs',multi_class="ovr")

model.fit(X_train,y_train)

model.intercept_ , model.coef_

```

Der folgende Graph zeigt die Entscheidungsgrenzen für das Klassifikationsmodell.
Alle Punkte innerhalb eines Bereiches werden der jeweiligen Klasse zugeordnet.

```python
i=0
xx = np.linspace(X_train[:,0].min()-1, X_train[:,0].max()+1, 300)
yy = np.linspace(X_train[:,1].min()-1, X_train[:,1].max()+1, 300)
XX, YY = np.meshgrid(xx,yy)
ZZ = model.predict(np.c_[XX.ravel(), YY.ravel()])
ZZ = ZZ.reshape(XX.shape)

yyy = model.intercept_[0] + model.coef_[0][0]*xx + model.coef_[0][1]*yy

colors = {0:'red', 1:'blue', 2:'green'}
markers = {0:'o', 1:'x', 2:'^'}
fig, ax = plt.subplots()
grouped = df.groupby('label')
#plt.pcolormesh(XX, YY, ZZ, cmap=plt.cm.Set3)
for key, group in grouped:
    group.plot(ax=ax, kind='scatter', x='x', y='y', label=key, marker=markers[key], color=colors[key])
    

plt.contour(XX, YY, ZZ, cmap=plt.cm.Blues)
plt.legend(loc='upper right', prop={'size': 12})
plt.savefig("LogistischeRegression21.png",transparent=True, dpi=300)
plt.show()
```